# Assignment 1.4 Pandas
Pandas is an open source Python package providing fast, flexible, and expressive data structures and data manipulation designed to make working with “relational” or “labeled” data both easy and intuitive. It is the fundamental high-level building block for doing practical, real world data analysis in Python.
https://pandas.pydata.org/

<!-- grade_id: cell-462d89bbf25b800a, points: 0, grade: false, solution: true -->
```python
## Import numpy and pandas

# YOUR CODE HERE
import pandas as pd
import numpy as np
```

<!-- grade_id: cell-ab515cd97836235b, points: 1, grade: true, solution: false -->
```python

```

The two primary data structures of pandas are 
- Series (1-dimensional) and 
- DataFrame (2-dimensional)

Read this [10 min intro to get started](https://pandas.pydata.org/pandas-docs/stable/10min.html)

### Data Loading

You are given two csv files.  
- 'sales1.csv'
- 'sales2.csv'

Load both CSV as dataframes and Investigate the data and its types

<!-- grade_id: cell-15bfe6e4718647be, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
sales1 = pd.read_csv("sales1.csv")
sales2 = pd.read_csv("sales2.csv")
```

Output:
```
Item Type  Order Date   Order ID   Ship Date  Units Sold  Unit Price  \
0   Cosmetics  10/18/2014  686800706  10/31/2014      8446.0      437.20   
1  Vegetables   11/7/2011  185941302   12/8/2011         NaN         NaN   
2   Baby Food  10/31/2016  246222341   12/9/2016      1517.0      255.28   
3     *Cereal   4/10/2010  161442649   5/12/2010      3322.0      205.70   
4      Fruits   8/16/2011  645713555   8/31/2011      9845.0        9.33   

   Unit Cost  
0     263.33  
1        NaN  
2     159.42  
3     117.11  
4       6.92
```

<!-- grade_id: cell-3c85ac1eab5f8b19, points: 0, grade: false, solution: true -->
```python
# Investigate the data 
# YOUR CODE HERE
print(sales1)
print(sales2)
print(sales1.describe())
print(sales2.describe())
```

Output:
```
Item Type  Order Date   Order ID   Ship Date  Units Sold  \
0          Cosmetics  10/18/2014  686800706  10/31/2014      8446.0   
1         Vegetables   11/7/2011  185941302   12/8/2011         NaN   
2          Baby Food  10/31/2016  246222341   12/9/2016      1517.0   
3            *Cereal   4/10/2010  161442649   5/12/2010      3322.0   
4             Fruits   8/16/2011  645713555   8/31/2011      9845.0   
5            *Cereal  11/24/2014  683458888  12/28/2014      9528.0   
6            *Cereal    3/4/2015  679414975   4/17/2015      2844.0   
7            Clothes   5/17/2012  208630645   6/28/2012      7299.0   
8         Vegetables   1/29/2015  266467225    3/7/2015      2428.0   
9         Vegetables  12/24/2013  118598544   1/19/2014      4800.0   
10           Clothes  12/29/2015  451010930   1/19/2016      3012.0   
11            Snacks   2/27/2010  220003211   3/18/2010      2694.0   
12         Household  11/17/2016  702186715  12/22/2016      1508.0   
13         Cosmet
```

### Merge Dataframes
Merge the two dataframes into a single dataframe "sales" based on order id such that new data consist of common order ID and all the other columns.
Also analyse the data in different columns.
Use [Merge function](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.merge.html)

<!-- grade_id: cell-82ee3339bf5f4570, points: 0, grade: false, solution: true -->
```python
sales = None
# YOUR CODE HERE
sales = pd.merge(sales1, sales2,on="Order ID", how="outer")
```

<!-- grade_id: cell-52d67c318540c515, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==686800706]["Country"].values == "Libya")
assert(sales[sales["Order ID"]==686800706]["Item Type"].values == "Cosmetics")
```

### Data Cleaning
If you look at the data we can see some issues,
- Some of the Item Type start with " * " we want to remove the asterik in the column
- Some columns data are missing indicated by NaN (u can try isnull(), isna() function)
  - If Region or Country data is missing replace it with unknown
  - if other data are missing, remove it from "sales" dataframe and add it in a "missing" dataframe 

Reference for this section:
- [Working with missing data](https://pandas.pydata.org/pandas-docs/stable/missing_data.html)
- [Python Lambda Expressions](https://docs.python.org/3/tutorial/controlflow.html#lambda-expressions)
- [Pandas Apply Function](https://pandas.pydata.org/pandas-docs/version/0.18/generated/pandas.Series.apply.html)
- [Pandas FillNa](https://pandas.pydata.org/pandas-docs/version/0.17.0/generated/pandas.DataFrame.fillna.html)

<!-- grade_id: cell-36c46b4d53bf011c, points: 0, grade: false, solution: true -->
```python
# Investigate dataframe for Null (NA/NAN) and count the number of rows with NA/NAN
NA_count = 0
# YOUR CODE HERE
NA_count = sales[sales.isnull().any(axis=1)].shape[0]

```

Output:
```
31
```

<!-- grade_id: cell-c3f7ee60f7aec5bf, points: 1, grade: true, solution: false -->
```python

```

<!-- grade_id: cell-04b3a523548b2603, points: 0, grade: false, solution: true -->
```python
# remove "*", NaN, add missing sales data to "missing" dataframe 
# YOUR CODE HERE
sales["Item Type"]=pd.Series(sales["Item Type"]).str.replace("*","",regex=True)
sales["Country"].fillna("unknown", inplace=True)
sales["Region"].fillna("unknown", inplace=True)
missing = pd.DataFrame(columns=sales.columns)
for index,row in sales.iterrows():
    if row.isna().sum():
        missing = missing.append(sales.loc[index, :], ignore_index=True)
        sales.drop(index, inplace=True)
```

<!-- grade_id: cell-7bc1b916b4472109, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==267066323]["Item Type"].values == "Cereal")
assert(sales[sales["Order ID"]==267066323]["Region"].values == "unknown")
```

Find the number of unique items available.

<!-- grade_id: cell-588bcfca26c28277, points: 0, grade: false, solution: true -->
```python
NumItems = 0
# YOUR CODE HERE
NumItems = len(pd.unique(sales["Item Type"]))
```

<!-- grade_id: cell-51b5cda4685cd3d0, points: 1, grade: true, solution: false -->
```python

```

### Grouping and Analysis 
Find out the statistical summary for ‘Unit Cost’ column for each Items.
Then extract the count, mean and std for "Cereal","Fruits","Vegetables","Baby Food" and "Meat" as a dataframe "Kitchen"
References:
- [Pandas Indexing](https://pandas.pydata.org/pandas-docs/stable/indexing.html)
- [Groupby](https://pandas.pydata.org/pandas-docs/stable/groupby.html)
- [Loc](https://pandas.pydata.org/pandas-docs/version/0.23.4/generated/pandas.DataFrame.loc.html)
- [filter](https://pandas.pydata.org/pandas-docs/stable/generated/pandas.DataFrame.filter.html)

<!-- grade_id: cell-d5b95b853fd96b60, points: 0, grade: false, solution: true -->
```python
Kitchen = None
# YOUR CODE HERE
indexdict = ["Cereal", "Fruits", "Vegetables", "Baby Food", "Meat"]
Kitchen = pd.DataFrame(index= indexdict, columns=["count", "mean", "std"])
for i in indexdict:
    unitCostCol = sales.loc[sales["Item Type"]== i ]["Unit Cost"]
    Kitchen.loc[ i ] = [unitCostCol.count(), unitCostCol.mean(), unitCostCol.std()]
```

<!-- grade_id: cell-986e659a273a364c, points: 1, grade: true, solution: false -->
```python
assert(Kitchen.loc["Cereal"]["count"]==77)
```

<!-- grade_id: cell-346f31cb4ea47c75, points: 0, grade: false, solution: true -->
```python
#Generate the new column ‘Total Profit’ for each orders (Total Profit = Total Revenue - Total Cost)
# YOUR CODE HERE
sales.insert(12, "Total Profit", value= sales["Total Revenue"] - sales["Total Cost"])
```

<!-- grade_id: cell-6419595bc71f6a09, points: 1, grade: true, solution: false -->
```python
assert(sales[sales["Order ID"]==246222341]["Total Profit"].values == 145419.62)
```

<!-- grade_id: cell-aae9d532c4dabf17, points: 0, grade: false, solution: true -->
```python
# How many data sets are there with profit greater than 60%
ans = None
# YOUR CODE HERE
sales["Total Profit%"] = sales["Total Profit"]/sales["Total Cost"]*100
ans = (sales["Total Profit%"]>=60).count()
```

<!-- grade_id: cell-c053463e75781632, points: 1, grade: true, solution: false -->
```python

```