# Assignment 1.1 Python
Python is an easy to learn, powerful programming language with efficient high-level data structures and object-oriented programming. Python’s elegant syntax and dynamic typing, together with its interpreted nature, make it an ideal language for scripting and rapid application development in many areas on most platforms.

### Jupyter Notebook
What you are reading now is an example of a Jupyter Notebook. The basic concept is that of a "notebook" containing text and programming code. You can easily edit the notebook using your web browser, run the programs in the ipython server in the background and see the output of the programs within the notebook. This is a powerful paradigm that is well suited to machine learning research, particularly when collaborating with other people.
For Reference:
- https://ipython.org/notebook.html
- https://jupyter.org/

```python
# This rectangular box is a "Cell" in Jupyter Notebook (click to edit)
# This line starting with # symbol is a comment
# Press run on the menu or CTRL + ENTER to run this cell
# or SHIFT + ENTER to run this cell and move to next cell
# to escape from the current cell press ESC and H to see shortcut keys available

var1 = 5
var2 = 6
print(var1 * var2)
```

Output:
```
30
```

Please go through https://docs.python.org/3/tutorial/introduction.html to get started with basic python introduction
Python also provides some built-in data types, in particular, dict, list, set and frozenset, and tuple. The str class is used to hold Unicode strings, and the bytes class is used to hold binary data. 
Visit the following page to know about the built-in data types
https://docs.python.org/3/library/stdtypes.html 

### TUPLES


```python
# Experiement here with your knowledge of Tuples

```

### Exercise 1
T is a Nested Tuple, use indexing to extract 'g' from the Tuple

<!-- grade_id: cell-3cc34b450273562e, points: 0, grade: false, solution: true -->
```python
T = ('a','b','c',('d','e'),('f',('f','g'),'k'),'p')
thisisg = None
# YOUR CODE HERE
thisisg = T[4][1][1]
```

<!-- grade_id: cell-b7541bf4e8c67052, points: 1, grade: true, solution: false -->
```python
assert(thisisg=='g')
### BEGIN HIDDEN hTESTS
assert(thisisg=='g')
```

### Exercise 2 (Lists)
Find the sum of the smallest 10 numbers of List L
Indexing, Sorted and sum function could be useful

<!-- grade_id: cell-31c3656a53f4497f, points: 0, grade: false, solution: true -->
```python
L =[75, 28, 47, 51, 96, 76, 32, 57, 10, 10, 25, 62, 17,  9, 68, 65, -2,
       54, 34, 41, 74, 83, 91, 59, -3, 88, 59, 91, 10, -9, 90, 46, 58, 25,
       62, 32, 74,  1, 61, 43, 25, 62, -5, 49, 40, 44, 83, -5, 13, 35]

# YOUR CODE HERE
ans = sum(sorted(L)[:10])
```

<!-- grade_id: cell-7dca8792586b992f, points: 1, grade: true, solution: false -->
```python
assert(ans==16)
```

## Control Flow
Python supports different Controls flow statements:
- While ...
- if ... elif ... else ...
- for ... in ...
- with ...

 https://docs.python.org/3/tutorial/controlflow.html

```python
# Experiment with Loops here; try printing out number which are divisible by 3
for i in range(1,100):
    if i%3==0:
        print(i)

```

Output:
```
3
6
9
12
15
18
21
24
27
30
33
36
39
42
45
48
51
54
57
60
63
66
69
72
75
78
81
84
87
90
93
96
99
```

### Exercise 3
Write a program which will list all such numbers which are divisible by 7 but are not a multiple of 5, between 2000 and 3200 (both included).logical Operator, for, if, and range could be helpful.


<!-- grade_id: cell-1b7aa023aac0ad93, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
for i in range(2000,3201):
    if (i%7==0 and i%5!=0):
        print(i)
```

Output:
```
2002
2009
2016
2023
2037
2044
2051
2058
2072
2079
2086
2093
2107
2114
2121
2128
2142
2149
2156
2163
2177
2184
2191
2198
2212
2219
2226
2233
2247
2254
2261
2268
2282
2289
2296
2303
2317
2324
2331
2338
2352
2359
2366
2373
2387
2394
2401
2408
2422
2429
2436
2443
2457
2464
2471
2478
2492
2499
2506
2513
2527
2534
2541
2548
2562
2569
2576
2583
2597
2604
2611
2618
2632
2639
2646
2653
2667
2674
2681
2688
2702
2709
2716
2723
2737
2744
2751
2758
2772
2779
2786
2793
2807
2814
2821
2828
2842
2849
2856
2863
2877
2884
2891
2898
2912
2919
2926
2933
2947
2954
2961
2968
2982
2989
2996
3003
3017
3024
3031
3038
3052
3059
3066
3073
3087
3094
3101
3108
3122
3129
3136
3143
3157
3164
3171
3178
3192
3199
```

<!-- grade_id: cell-e4a62a5f25d70f6b, points: 1, grade: true, solution: false -->
```python

```

### Exercise 4 (Function) 
write a function fibbo to return a list of "n" (greater than 2) fibonnaci number

<!-- grade_id: cell-999e97f969ad859a, points: 0, grade: false, solution: true -->
```python
def fibbo(n):
# YOUR CODE HERE
    fiblist = [1, 1]
    a = 1
    b = 1
    for i in range(n-2):
        c = a + b
        a = b
        b = c
        fiblist.append(b)
    return fiblist
```

<!-- grade_id: cell-f97bd7a6d835fcb9, points: 1, grade: true, solution: false -->
```python
assert(fibbo(5)==[1, 1, 2, 3, 5])
```

### Exercise 5 (Strings)
Go through the documentation of string to work through the following section
https://docs.python.org/3/library/string.
Combine the following strings and list (as comma seperated list) into a single string.

<!-- grade_id: cell-647dfc2b970c2bea, points: 0, grade: false, solution: true -->
```python
str1 = "Welcome to Fuse.ai. "
str2 = "This is the First module. "
str3 = "We will cover "
modules_list = ["Python", "Python Classes", "Numpy", "Pandas", "Matplotlib", "etc."]
combinedstr = None
# YOUR CODE HERE
combinedstr = str1 + str2 + str3 + (", ").join(modules_list)

```

<!-- grade_id: cell-07b4653968423c7b, points: 1, grade: true, solution: false -->
```python
assert(combinedstr=='Welcome to Fuse.ai. This is the First module. We will cover Python, Python Classes, Numpy, Pandas, Matplotlib, etc.')
```

### Exercise 6 (Set)
find the common elements of a and b and not in c

<!-- grade_id: cell-658c09dc97713ac7, points: 0, grade: false, solution: true -->
```python
a = range(10,27)
b = [1,4,5,8,2,1,9,12,79,35,13,24]
c = (3,7,12,90,45)

def common(a,b,c):
# YOUR CODE HERE
    return set(a) & set(b) - set(c)
```

<!-- grade_id: cell-e0d4e0cb5111b294, points: 1, grade: true, solution: false -->
```python
assert(common(a,b,c)=={13,24})
```

### Exercise 7 (Dictionaries)
Ram went to buy some vegetables in Vegetable market. Write a function to find his total cost given rams shopping list and price dictionary (price are in Rupees)


<!-- grade_id: cell-9e2bdb0f9fe2202a, points: 0, grade: false, solution: true -->
```python
price = {"Tomato Big":50,"Tomato Small":90,"Potato Red":70,"Onion Dry":45,"Carrot":120,"Cabbage":40,"Cauli Local":60,"Raddish White":40,"Brinjal Long":40,"Brinjal Round":50,"Cow pea":90,"French Bean":60,"Soyabean Green":80,"Bitter Gourd":90,"Bottle Gourd":40,"Pointed Gourd":50,"Smooth Gourd":60,"Pumpkin":40,"Squash":60,"Okara":90,"Barela":60,"Arum":60,"Christophine":20,"Brd Leaf Mustard":60,"Spinach Leaf":100,"Cress Leaf":100,"Mustard Leaf":60,"Fenugreek Leaf":100,"Onion Green":110,"Mushroom":180,"Neuro":80,"Sugarbeet":90,"Lettuce":60,"Celery":180,"Parseley":180,"Fennel Leaf":100,"Mint":180,"Turnip A":60,"Tamarind":150,"Bamboo Shoot":80,"Tofu":100,"Gundruk":220,"Apple":120,"Banana":80,"Lime":250,"Pomegranate":200,"Orange":130,"Water Melon":60,"Sweet Orange":120,"Pineapple":120,"Cucumber":70,"Papaya":80,"Guava":60,"Mombin":50,"Ginger":130,"Chilli Dry":250,"Chilli Green":90,"Capsicum":90,"Garlic Green":290,"Coriander Green":100,"Garlic Dry Chinese":150,"Garlic Dry Nepali":130,"Fish Fresh":270}

def totalcost(rams_list):
# YOUR CODE HERE
    total = 0
    for i in rams_list:
        total+=price[i]
    return total
```

<!-- grade_id: cell-c19b7cf0c33943c5, points: 1, grade: true, solution: false -->
```python
assert(totalcost(["Tomato Big","Tomato Small"])==140)
```

### Exercise 8 (Date Time)
Python supports datetime module which supplies functons and classes for manipulating dates and times in both simple and complex ways.
https://docs.python.org/3/library/datetime.html

Import datetime

<!-- grade_id: cell-427eed3d24974b01, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
import datetime
```

<!-- grade_id: cell-725b73afe2364885, points: 1, grade: true, solution: false -->
```python

```

### Exercise 9
Since we all love holidays. 
Write a function to test if the day of the week is saturday or sunday

<!-- grade_id: cell-1eff2257a92fc3ca, points: 0, grade: false, solution: true -->
```python
def isholiday(date):
# YOUR CODE HERE
    return date.weekday()==6 or date.weekday()==5
```

<!-- grade_id: cell-4bf249a294e13b37, points: 1, grade: true, solution: false -->
```python
assert(isholiday(datetime.datetime(2018, 10, 21))==True)
```

### List Comprehension
https://docs.python.org/3/tutorial/datastructures.html#list-comprehensions

Also go through 
- https://docs.python.org/3/library/functions.html#zip 
- https://docs.python.org/3/library/itertools.html, 

as these functions are very useful in different situation

```python
# Experiement with List comprehension here

```

### Exercise 10
Carryout matrix multiplication using list comprehension. Zip function might come handy for this task

<!-- grade_id: cell-03919f8e4b02f4c6, points: 0, grade: false, solution: true -->
```python
def matmultiply(A, B):
    # Write Solution in 1 line
# YOUR CODE HERE
    return [[sum(x*y for x,y in zip(row_A,col_B)) for col_B in zip(*B)] for row_A in A]
```

<!-- grade_id: cell-156372dc6fe33108, points: 1, grade: true, solution: false -->
```python

```